---
title_meta  : Chapter 3
title       : Functions and packages
description : Add chapter description here.

--- type:VideoExercise lang:python xp:50 skills:2
## Functions

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

--- type:NormalExercise lang:python xp:100 skills:2
## Familiar functions

Out of the box, Python offers a bunch of built-in functions to make your life as a data scientist easier. You already know two such functions: `print()` and `type()`. Also the functions `str()`, `int()`, `bool()` and `float()` to switch between data types are built-in functions.

Calling a function is easy. To get the type of `3.0` and store the output as a new variable, `result`, you can use the following:

```
the_type = type(3.0)
```

The general 'recipe' for calling functions is thus:

```
result = function_name(input)
```

In this and the following exercise, you'll be working with ... ?

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
# sct code
```



--- type:VideoExercise lang:python xp:50 skills:2
## Methods

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

--- type:VideoExercise lang:python xp:50 skills:2
## Packages

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

