---
title_meta  : Chapter 4
title       : Numpy
description : Add chapter description here.

--- type:VideoExercise lang:python xp:50 skills:2
## Numpy: from coding to math

Introduce numpy. Very import feature: numpy.array. Convert a normal list into a mathematical object. Vector arithmetic. What does + do to a normal list: concat.

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

--- type:NormalExercise lang:python xp:100 skills:2
## Getting started with numpy

This is an easy one. Just import numpy. As we are going to use numpy a lot, it's best to import it as np.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:NormalExercise lang:python xp:100 skills:2
## Calculate the total

Given: budget spent on dinner last week. budget spent on dinner this week. Add them up. You want to calculate the price in euros: multiply.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:NormalExercise lang:python xp:100 skills:2
## Tipping to survive

Being a generous person, you always tip for service in a restaurant. However, the mood you're in will greatly influence the percentage of the dinner price you'll tip. On mondays you are in a bad mood, you only tip 5%. On tuesdays till thursday you'll tip 15%. In the weekends you're always in a great mood, so you'll tip 25%. Use numpy (np.array(dinner_prices)*(np.array([0.05, 0.15, 0.15, 0.15, 0.25, 0.25, 0.25]) + 1))

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```




