---
title_meta  : Chapter 1
title       : The first chapter of the introduction to Python course
description : Description of the first chapter of the introduction to Python course

--- type:NormalExercise xp:100 lang:python
## First Normal Exercise Title
  
This is the first Python exercise.

*** =instructions
- Assign the value 5 to the variable `x`.
- Print the variable `x`.

*** =hint
This is a hint.

*** =pre_exercise_code
```{python}
y = 7
```

*** =sample_code
```{python}
# Assign the value 5 to the variable x


# Print the variable x

```

*** =solution
```{python}
# Assign the value 5 to the variable x
x = 5

# Print the variable x
print(x)
```

*** =sct
```{python}
# SCT CODE COMES HERE...
```

*** =skills
2

--- type:NormalExercise xp:100 lang:python
## Another normal exercise
  
Just random shtuff.

*** =instructions
- Import the `matplotlib` package
- Print the variable `x`.
- Another instruction

*** =hint
This is a hint.

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}
# import matplotlib


# Make a plot

```

*** =solution
```{python}
# import matplotlib
import matplotlib.pyplot as plt

# Make a plot
plt.plot([1,2,3,4])
plt.ylabel('some numbers')
plt.show()
```

*** =sct
```{python}
# SCT CODE COMES HERE...
```

*** =skills
2

--- type:MultipleChoiceExercise xp:50 lang:python
## First Multiple Choice Exercise Title

This is a python multiple choice exericse.

*** =instructions
- Option 1
- Option 2

*** =hint
The Hint

*** =pre_exercise_code
```{python}
# none
```

*** =sct
```{python}

```

*** =skills
2
