---
title_meta  : Chapter 1
title       : Python Basics
description : An introduction to the basic concepts of Python and its place in the Data Science community

--- type:VideoExercise lang:python xp:50 skills:2
## First steps into Python

A brief history of Python. It's place in data science. The pros and cons of using it in data science. Using the Python console for basic calculations.

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

--- type:MultipleChoiceExercise lang:python xp:50 skills:2
## When to use it

When to use Python? A few use cases are given. You want to do a quick calculation. You want to build a robust software package. You want to analyse some data. All of the above. 

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sct
```{python}
test_mc(2) # if 2 is the correct option.
```

--- type:NormalExercise lang:python xp:100 skills:2
## Crunch those numbers

You and your friends are splitting the tab. You drank 3 cocktails and 5 beers in total. You want to split the tab evenly. 

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:NormalExercise lang:python xp:100 skills:2
## Generousness comes a long way

Instead of splitting the tab evenly you want to pay 50% of the tab yourself and let your friends pay the rest. Your friend think the same thing and pay 50% of what remains (use powers)...

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:VideoExercise lang:python xp:50 skills:2
## Variables and their types

Introducing variables and variable assignment. The difference between types of variables (using strings and booleans).

*** =video_link
```{python}
//player.vimeo.com/video/108225030
```

--- type:MultipleChoiceExercise lang:python xp:50 skills:2
## Which is what

A list of variables is given with a value. Choose the choice which correctly lists their types.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sct
```{python}
test_mc(2) # if 2 is the correct option.
```

--- type:NormalExercise lang:python xp:100 skills:2
## Assign all the variables

Instead of remembering the cost of the tab you decide to assign it to a variable for later use. There is one new person whos name you can't seem to remember so you decide to decide his name to a variable as well. Finally you decide to assign a boolean variable to remember the fact whether you were drunk or not.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:NormalExercise lang:python xp:100 skills:2
## Calculations with variables

You wake up the next morning with the variables still in your memory. You decide to calculate how much you should pay Josh who paid the whole tab. You decide to give him somewhat extra based on that amount.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:NormalExercise lang:python xp:100 skills:2
## Invalid calculations

You find some other variables in your memory of which you have little memory. You decide to check out what their type is. After finding out their type you decide to experiment with them a little bit.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


--- type:MultipleChoiceExercise lang:python xp:50 skills:2
## Which are valid and which aren't

You are given a list of variables and operations on them. Select which one will result in an error.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sct
```{python}
test_mc(2) # if 2 is the correct option.
```

--- type:NormalExercise lang:python xp:100 skills:2
## Did we pay everything? 

Given the total amount of the tab. Also given - an amount that everyone wants to pay. Do they sum up to the total amount?

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sample_code
```{python}
# sample code comes here
```

*** =solution
```{python}
# Solution code
```

*** =sct
```{python}
# sct code comes here
```


