---
title_meta  : Chapter 2
title       : Lists
description : Add chapter description here.

--- type:VideoExercise lang:python xp:50 skills:2
## Lists, what are they?

*** =video_link
//player.vimeo.com/video/108225030

--- type:NormalExercise lang:python xp:100 skills:2
## Create a list

As opposed to `int`, `bool` etc, a list is a _compound data type_: you can create a collection of values.

To create a list you use square brackets and commas, like this:

```
my_list = ["my", "first", "list"]
```

You can also use variables you've created before:

```
a = 2
b = 4
my_list = [a, b]
```

After measuring the height of your family, you decide to collect some information on the house you're living in. The areas of the different parts of your crib are stored in separate variables for now, as shown in the script.

*** =instructions
- Create a list, `areas`, that contains the area of the hallway, kitchen, living room, bedroom and bathroom.
- Print `areas` with the `print()` function.

*** =hint
- You can use the variables that have already been created to build the list: `areas = [hallway, kitchen, ...]`.
- Type `print(areas)` in your script to print out the list when submitting.

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
# sct code
```

--- type:NormalExercise lang:python xp:100 skills:2
## Create list with different types

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
# sct code
```

--- type:NormalExercise lang:python xp:100 skills:2
## Select the valid list

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
# sct code
```

--- type:NormalExercise lang:python xp:100 skills:2
## List of lists

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec
```

*** =sample_code
```{python}
# sample code
```

*** =solution
```{python}
# solution code
```

*** =sct
```{python}
# sct code
```

