---
title_meta  : Chapter 2
title       : Lists
description : This a chapter about lists

--- type:MultipleChoiceExercise lang:python xp:50 skills:2
## My MultipleChoiceExercise

Assignment comes here. Use Markdown for text formatting.

*** =instructions
- option 1
- option 2
- option 3

*** =hint
hint comes here

*** =pre_exercise_code
```{python}
# pec comes here
```

*** =sct
```{python}
test_mc(2) # if 2 is the correct option.
```

